Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **37 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Adaobi Amanna"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Asymptotic notations; maximum-subarray problem

## Question 1 (Exercise 3.1-3 of Cormen, et al. ) [time estimate: 5 minutes]
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is true but meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
The Big O notation tells us about the asymptotic behavior of an Algorithm A. This means that it majorly characterizes the upperbound of the algorithm. However, in this case, it is saying that the running time is 'at least' O(n^2), which is not completely useless but it does not tell us entirely what we are supposed to know about the upper bound, it leaves it open ended because the behavior of the upper bound can exceed O(n^2); basically, it defines the lower bound. Hence it is not meaningful as it fails to provide the needed information about the behavior of the upper bound. 

## Question 2 (Exercise 3.1-4 of Cormen, et al. ) [time estimate: 5 minutes]

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$? Why?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
 $2^{n+1}=O(2^n)$ is true because 1 can be viewed as a constant that does not significantly affect the scaling of the algorithm in the long run. The behavior as n increases will still be similar, hence we can get rid of the constant and still perfectly express the behavior of the algorithm.
 
 However, $2^{2n}=O(2^n)$ is false because, unlike the first case where the addition of 1 is a constant, here we are dealing with an exponential growth where the behavior of $2^{2n}$ will differ significantly from the behavior of $O(2^n)$. With $2^{2n}$, we are squaring the value of $2^{n}$ by raising it by 2. On the other hand, $O(2^n)$ is only raised to the power of n, and the scaling of the power 2n is definitely going to be bigger than just n.

## Question 3 [time estimate: 20 minutes]
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    #initialize the start index, the stop index and the value fo maximum array
    start_idx = 0
    stop_idx = 0
    Max_arr = float('-inf')
    
    #loop through the given array
    for i in range(len(A)-1):
        #run another loop for values in the range after i
        for j in range(i+1, len(A)):
            #sum the elements in the array interval
            total = sum(A[i:j])
            #evaluate if the sum is bigger than maximum array sum
            if total > Max_arr:
                #If true, update the parameters and loop again
                Max_arr = total
                start_idx = i
                stop_idx = j-1
    return (start_idx, stop_idx, Max_arr)
print(bruteforce_max_subarray([13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]) == (7,10,43))
```

Output:
```
True
```

```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

## Question 4 [time estimate: 2 minutes]
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
print(bruteforce_max_subarray([13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]))
```

Output:
```
(7, 10, 43)
```

## Question 5. Asymptotic notation [time estimate: 5 minutes]
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Running time 
--- | --- 
Insertion sort | 
Selection sort |  
Bubble sort |  
Finding maximum subarray |  

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
Algorithm | Running time 
--- | --- 
Insertion sort | $O(n)$
Selection sort | $\\Theta(n^2)$  
Bubble sort | $\\Theta(n^2)$ 
Finding maximum subarray | $\\Theta(n^2)$

## [Optional] Question 6 [time estimate: 15 minutes] 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """
    Implements brute-force minimum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the minimum subarray
        - the end index of the minimum subarray
        - the value of the minimum subarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```